# Week1 Assignment - Testing OKPy and the autograder.

For this assignment, you will load the iris flowers dataset from sklearn and predict flower species based on their petal and sepal sizes.

```python
"""This is a docstring!"""
```

Output:
```
'This is a docstring!'
```

$\beta=5$

```python
alpha = 4
```

```python
%who
```

Output:
```
alpha
```

```python
var_list = %who_ls
```

```python
var_list
```

Output:
```
['alpha']
```

## Part A. Load the Iris dataset.

Sklearn comes with prepared datasets that you can load. To load the iris dataset, run:
```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
iris.target_names
```

Output:
```
array(['setosa', 'versicolor', 'virginica'], dtype='<U10')
```

```python
iris.target
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets

# Load the iris dataset
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-9-eb43bd51aa77> in <module>
      3 # Load the iris dataset
      4 # YOUR CODE HERE
----> 5 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert (iris['data'] == datasets.load_iris()['data']).all()
```

## Part B. Use an SVM to predict iris species.

To setup the data and train an SVM, run the following:
```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf")
classifier.fit(X, y)
predictions = classifier.predict(X)
```

```python
%who
```

```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC()

X.shape

classifier.fit(X, y)

predictions = classifier.predict(X)
```

Output:
```
/srv/conda/envs/notebook/lib/python3.7/site-packages/sklearn/svm/base.py:193: FutureWarning: The default value of gamma will change from 'auto' to 'scale' in version 0.22 to account better for unscaled features. Set gamma explicitly to 'auto' or 'scale' to avoid this warning.
  "avoid this warning.", FutureWarning)
```

```python
from sklearn.datasets import load_iris
```

```python
i = load_iris()
```

```python
type(i)
```

- Overfitting (Variance): tailored to a specific dataset (does not generalize well), misses the point, memorizes the data
- Underfitting (Bias): accuracy could be better, too general 

1. Import
2. Instantiate (a class)
  - Estimators:
    - Classifiers
    - Regressors
3. Fit (fit())
4. Predict (predict())
5. Evaluate (score())


```python
predictions
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

```python
from sklearn.model_selection import train_test_split
```

```python
classifier.score(X,y)
```

Output:
```
0.9866666666666667
```

```python
from sklearn.svm import SVC
```

```python
from sklearn.svm import SVC

X = iris.data
y = iris.target
```

```python
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.33, random_state=42)
```

```python
iris.target
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

```python
classifier = SVC()
```

```python
X.shape
```

Output:
```
(150, 4)
```

```python
X_train.shape
```

Output:
```
(100, 4)
```

```python
y.shape
```

Output:
```
(150,)
```

```python
y_train.shape
```

Output:
```
(100,)
```

```python
y_test.shape
```

Output:
```
(50,)
```

```python
X_test.shape
```

Output:
```
(50, 4)
```

```python
classifier
```

Output:
```
SVC(C=1.0, cache_size=200, class_weight=None, coef0=0.0,
    decision_function_shape='ovr', degree=3, gamma='auto_deprecated',
    kernel='rbf', max_iter=-1, probability=False, random_state=None,
    shrinking=True, tol=0.001, verbose=False)
```

```python
classifier.fit(X_train, y_train)
```

Output:
```
/srv/conda/envs/notebook/lib/python3.7/site-packages/sklearn/svm/base.py:193: FutureWarning: The default value of gamma will change from 'auto' to 'scale' in version 0.22 to account better for unscaled features. Set gamma explicitly to 'auto' or 'scale' to avoid this warning.
  "avoid this warning.", FutureWarning)
SVC(C=1.0, cache_size=200, class_weight=None, coef0=0.0,
    decision_function_shape='ovr', degree=3, gamma='auto_deprecated',
    kernel='rbf', max_iter=-1, probability=False, random_state=None,
    shrinking=True, tol=0.001, verbose=False)
```

```python
classifier.score(X_test,y_test)
```

Output:
```
1.0
```

```python

```

```python

```

```python
from sklearn.metrics 
```

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-13-0bb9f2249024> in <module>
----> 1 claasifier.score(X,y)
      2 
      3 from sklearn.svm import SVC
      4 
      5 # YOUR CODE HERE

NameError: name 'claasifier' is not defined
```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150, 4)
assert len(y) == 150
```